---
title_meta  : Chương 4
title       : Factors
description : Trong rất nhiều trường hợp, dữ liệu có giá trị nằm trong một nhóm hữu hạn các giá trị cho trước. Ví dụ, giới tính có giá trị nam hoặc nữ. Trong R, các biến có loại này được gọi là factor. Các biến factor đóng vai trò rất quan trọng trong quá trình phân tích dữ liệu. Do đó, chúng ta sẽ cùng học cách tạo và xử lý các biến factor trong bài giảng dưới đây.

--- type:NormalExercise xp:100 skills:1 key:05273321916d99bb9c0deadf75c6834d25a47244
## Factor là gì và tại sao chúng ta lại sử dụng chúng?

Trong chương này, chúng ta sẽ cùng nhau tìm hiểu các biến **factors**.

Thuật ngữ factor được sử dụng cho loại dữ liệu dùng trong phân tích thống kê được sử dụng để lưu trữ các nhóm. Sự khác biệt cơ bản giữa biến nhóm và biến số là ở chỗ, các giá trị của biến nhóm thuộc một **nhóm hữu hạn các giá trị cho trước**. Trong khi đó, biến định dạng số có thể có vô số các giá trị. (Trong một số tài liệu tiếng Việt, còn được gọi là **biến yếu tố**).

Khi phân tích dữ liệu, việc xác định được đang phân tích biến số hay biến nhóm (factor) là rất quan trọng, bởi lẽ các mô hình phân tích thống kê sẽ xử lý hai định dạng dữ liệu này khác nhau (Các bạn sẽ thấy rõ điều nay trong các phần tiếp theo).

Một ví dụ điển hình của biến nhóm là biến 'Giới tính'. Mỗi cá nhân có thể có giới tính là "Nam" hoặc "Nữ". Do đó, "Nam" và "Nữ" là hai giá trị của biến định dạng nhóm "Giới tính", và bất kỳ quan sát nào cũng có thể gán một trong hai giá trị - "Nam" hoặc "Nữ"

*** =instructions
Gán biến `theory` giá trị của `factors for categorical variables"`.

*** =hint
Sử dụng ký tự gán như thông thường (`<-`); 

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Gán nội dung của chương này cho biến theory
```

*** =solution
```{r}
# Gán nội dung của chương này cho biến theory
theory <- "factors for categorical variablesm"
```

*** =sct
```{r}
test_object("theory", incorrect_msg = "Kiểm tra lại chuỗi giá trị `\"factors for categorical variables\"` gán cho biến `theory`. Lưu ý rằng R có phân biệt chữ hoa và chữ thường.");
success_msg("Rất tốt! Bạn đã sẵn sàng chưa? Tiếp tục bài tập tiếp theo nhé!")
```


--- type:NormalExercise xp:100 skills:1 key:6cc21c842b075347926bb1b244782213df32e370
## Factor là gì và tại sao chúng ta lại sử dụng chúng? (2)

Để tạo biến factor trong R, bạn có thể sử dụng hàm [`factor()`](http://www.rdocumentation.org/packages/base/functions/factor). Đầu tiên, bạn cần phải tạo một véc-tơ chứa tất cả các quan sát thuộc một nhóm các giá trị hữu hạn. Ví dụ, `gender_vector`chưá 5 quan sát độc lập:


```
gender_vector <- c("Male","Female","Female","Male","Male")
```

Trong trường hợp này, véc-tơ có hai giá trị: "Male" và "Female" (trong R còn được gọi là **'factor levels'**).

Hàm [`factor()`](http://www.rdocumentation.org/packages/base/functions/factor) sẽ chuyển đổi giá trị của véc-tơ thành dạng factor:

```
factor_gender_vector <- factor(gender_vector)
```

*** =instructions
- Chuyển véc-tơ dạng character `gender_vector` thành véc-tơ có dạng `factor()` và gán kết quả cho véc-tơ `factor_gender_vector`
- Hiển thị giá trị của véc-tơ `factor_gender_vector` và chắc chắc rằng R hiển thị các giá trị của factor bên dưới khu vực hiển thị giá trị của véc-tơ

*** =hint
Sử dụng hàm [`factor()`](http://www.rdocumentation.org/packages/base/functions/factor) với véc-tơ `gender_vector`. Xem ví dụ và tài liệu đi kèm.
*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Tạo gender_vector
gender_vector <- c("Male", "Female", "Female", "Male", "Male")

# Chuyển đổi gender_vector thành dạng factor
factor_gender_vector <-

# Hiển thị véc-tơ factor_gender_vector

```

*** =solution
```{r}
# Tạo gender_vector
gender_vector <- c("Male", "Female", "Female", "Male", "Male")

# Chuyển đổi gender_vector thành dạng factor
factor_gender_vector <- factor(gender_vector)

# Hiển thị véc-tơ factor_gender_vector
factor_gender_vector
```

*** =sct
```{r}
test_object("factor_gender_vector",
            incorrect_msg = "Bạn đã gán gán factor `gender_vector` cho véc-tơ `factor_gender_vector` chưa?")
test_output_contains("factor_gender_vector", incorrect_msg = "Nhớ hiển thị véc-tơ `factor_gender_vector`!")
success_msg("Rất tốt! Nếu bạn muốn tìm hiểu thêm về hàm `factor()`, gõ lệnh `?factor` trên màn hình console. Lệnh này sẽ mở trang tài liệu chứa các thông tin về hàm 'factor'. Tiếp tục làm bài tập tiếp theo nhé.");
```


--- type:NormalExercise xp:100 skills:1 key:5bd4f50afc2c2dbc881e16b8ca94ca56960dff42
## Factor là gì và tại sao chúng ta lại sử dụng chúng? (3)

Biến yếu tố (nhóm) có 2 loại: **biến yếu tố không sắp xếp thứ tự** và **biến yếu tố có sắp xếp thứ tự** (nominal & cardinal variable)

Đối với biến yếu tố không sắp xếp thứ tự (nominal variable), ta không thể nói rằng "giá trị này tốt hơn giá trị kia". Ví dụ, khi sử dụng véc-tơ `animals_vector` với các giá trị là `"Elephant"`,`"Giraffe"`, `"Donkey"` và `"Horse"`, ta không thể nói loài nào tốt hơn hay kém hơn loài nào (Hẳn một vài bạn sẽ không đồng ý ;-) ).

Ngược lại, biến yếu tố có sắp xếp thứ tự cho phép sắp xếp các giá trị của biến theo thứ tự. Ví dụ, đối với véc-tơ `temperature_vector` với các giá trị: `"Low"`, `"Medium"` và `"High"`, giá trị `"High"` có giá trị lớn hơn `"Medium"` và  `"Medium"` có giá trị lớn hơn `"Low"`.

*** =instructions

Ấn 'Submit Answer' để kiểm tra cách R triển khai và hiển thị các  biến yếu tố có sắp xếp và không sắp xếp thứ tự. Đừng lo nếu bạn chưa hiểu hết các đoạn code, chúng ta sẽ học trong các phần tiếp theo.

*** =hint
Ấn nút 'Submit Answer' và xem màn hình hiển thị. Lưu ý cách thức R hiển thị thứ tự của các biến yếu tố có sắp xếp thứ tự (ordinal)

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Các loài động vật
animals_vector <- c("Elephant", "Giraffe", "Donkey", "Horse")
factor_animals_vector <- factor(animals_vector)
factor_animals_vector

# Nhiệt độ
temperature_vector <- c("High", "Low", "High","Low", "Medium")
factor_temperature_vector <- factor(temperature_vector, order = TRUE, levels = c("Low", "Medium", "High"))
factor_temperature_vector
```

*** =solution
```{r}
# Các loài động vật
animals_vector <- c("Elephant", "Giraffe", "Donkey", "Horse")
factor_animals_vector <- factor(animals_vector)
factor_animals_vector

# Nhiệt độ
temperature_vector <- c("High", "Low", "High","Low", "Medium")
factor_temperature_vector <- factor(temperature_vector, order = TRUE, levels = c("Low", "Medium", "High"))
factor_temperature_vector
```

*** =sct
```{r}
msg <- "Không thay đổi bất kỳ điều gì ở đoạn code mẫu. Ấn nút Submit Answer để xem lời giải"
test_object("animals_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("temperature_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_animals_vector", undefined_msg = msg, incorrect_msg = msg)
test_output_contains("factor_animals_vector", incorrect_msg = msg)
test_object("factor_temperature_vector", undefined_msg = msg, incorrect_msg = msg)
test_output_contains("factor_temperature_vector", incorrect_msg = msg)
success_msg("Bạn đã nắm được kiến thức trong bài tập này? Tuyệt vời! Hãy tiếp tục bài tập tiếp theo để hiểu rõ hơn về factor levels")
```


--- type:NormalExercise xp:100 skills:1 key:1aa698978d32d1a0befa4700d7da85a648e1d69e
## Các giá trị của factor (factor levels) 

Khi mới tiếp cận dữ liêu, các bạn hẳn sẽ nhận thấy rằng các factor chứa một số giá trị nhất định. Tuy nhiên, đôi khi ta muốn đổi tên của các giá trị này nhằm phục vụ quá trình phân tích dữ liệu. R cho phép bạn thực hiện quá trình này với hàm [`levels()`](http://www.rdocumentation.org/packages/base/functions/levels):

```
levels(factor_vector) <- c("name1", "name2",...)
```


Một ví dụ điển hình để giải thích quá trình này là sử dụng bảng hỏi. Khi thu thập ý kiến, một câu hỏi điển hình là giới tính của người được hỏi. Hẳn các bạn vẫn còn nhớ đây là biến factor có hai giá `"M"` và `"F"`.

```
survey_vector <- c("M", "F", "F", "M", "M")
```

Tiếp đó, khi bắt đầu phân tích dữ liệu, quan tâm chính của bạn sẽ là giữ được cái nhìn tổng quan của tất cả các biến cũng như ý nghĩa của chúng. Khi đó, để giúp quá trình phân tích được dễ dàng hơn, bạn thường sẽ muốn đổi các giá trị của factor thành `"Male"` và `"Female"` (nam và nữ) thay cho giá trị `"M"` và `"F"`

**Lưu ý**: Thứ tự của các giá trị trong biến factor rất quan trọng trong quá trình phân tích. Nếu gõ `levels(factor_survey_vector)`, bạn sẽ thấy kết quả hiển thị là `[1] "F" "M"`. Nếu bạn không xác định rõ giá trị của facotr, `R` sẽ tự động gán chúng theo thứ tự bảng chữ cái. Để gán `"F"` thành `"Female"` và `"M"` thành `"Male"`, giá trị của factor cần được gán thành `c("Female", "Male")` theo thứ tự

*** =instructions
- Kiểm tra đoạn code xây dựng giá trị của factor từ `survey_vector`. Bạn nên sử dụng `factor_survey_vector` trong phần bài giảng tiếp theo.
- Đổi thứ thự các giá trị của `factor_survey_vector` thành `c("Female", "Male")`. Lưu ý thứ tự các giá trị trong vector


*** =hint
Lưu ý thứ tự các giá trị trong factor. Bạn có thể kiểm tra thứ tự của giá trị trong factor bằng cách sử dụng cầu lệnh `levels(factor_survey_vector)`.


*** =pre_exercise_code
```{r}
# no pec
survey_vector <- c("M", "F", "F", "M", "M")
factor_survey_vector <- factor(survey_vector)
```

*** =sample_code
```{r}
# Xây dựng véc-tơ factor_survey_vector
survey_vector <- c("M", "F", "F", "M", "M")
factor_survey_vector <- factor(survey_vector)

# Xác định thứ tự của factor_survey_vector
levels(factor_survey_vector) <-

factor_survey_vector
```

*** =solution
```{r}
# Xây dựng véc-tơ factor_survey_vector
survey_vector <- c("M", "F", "F", "M", "M")
factor_survey_vector <- factor(survey_vector)

# Xác định thứ tự của factor_survey_vector
levels(factor_survey_vector) <- c("Female", "Male")

factor_survey_vector
```

*** =sct
```{r}
msg = "Đừng thay đổi định nghĩa của véc-tơ `survey_vector`!"
test_object("survey_vector", undefined_msg = msg, incorrect_msg = msg)
msg = "Không chỉnh sửa các câu lệnh tạo factor"
test_function("factor", "x", not_called_msg = msg, incorrect_msg = msg)
test_object("factor_survey_vector", eq_condition = "equal",
            incorrect_msg = paste("Bạn có gán giá trị chính xác của factor cho `factor_survey_vector` chưa? Sử dụng hàm `levels(factor_survey_vector) <- c(\"Female\", \"Male\")`. Nhớ rằng R có phân biệt chữ in hoa và chữ in thường nhé"))
success_msg("Tuyệt vời! Tiếp tục làm bài tập tiếp theo nhé.")
```


--- type:NormalExercise xp:100 skills:1 key:a549f13c0644ccc89cd39a10aa48706754637ed0
## Tóm tắt giá trị của factor


Sau khi kết thúc khóa học này một trong những câu lệnh bạn sẽ hay dùng nhất là [`summary()`](http://www.rdocumentation.org/packages/base/functions/summary). Câu lệnh này cho phép bạn nhìn được tổng quan các biến trong tập dữ liệu của mình:

```
summary(my_var)
```

Quay lại với ví dụ về bảng hỏi ở phần trước, bạn có lẽ sẽ muốn kiểm tra xem có bao nhiêu giá trị `"Male"` và `"Female"`. Hàm [`summary()`](http://www.rdocumentation.org/packages/base/functions/summary) sẽ trả lời cho câu hỏi trên của bạn. 

*** =instructions
Sử dụng hàm [`summary()`](http://www.rdocumentation.org/packages/base/functions/summary) đối với `survey_vector` và `factor_survey_vector`. Interpret the results of both vectors. Are they both equally useful in this case? Giả thích kết quả của cả 2 véc-tơ. Liệu rằngi kết quả đều có giá trị như nhau trong trường hợp trên? 

*** =hint
Gọi hàm [`summary()`](http://www.rdocumentation.org/packages/base/functions/summary) đối với  `survey_vector` và `factor_survey_vector`, đơn giản vậy thôi!

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Tạo véc-tơ factor_survey_vector 
survey_vector <- c("M", "F", "F", "M", "M")
factor_survey_vector <- factor(survey_vector)
levels(factor_survey_vector) <- c("Female", "Male")
factor_survey_vector

# Tổng hợp kết quả từ véc-tơ survey_vector


# Gọi kết quả tổng hợp từ véc-tơ  factor_survey_vector

```

*** =solution
```{r}
# Tạo véc-tơ factor_survey_vector 
survey_vector <- c("M", "F", "F", "M", "M")
factor_survey_vector <- factor(survey_vector)
levels(factor_survey_vector) <- c("Female", "Male")
factor_survey_vector

# Tổng hợp kết quả từ véc-tơ survey_vector
summary(survey_vector)

# Gọi kết quả tổng hợp từ véc-tơ  factor_survey_vector
summary(factor_survey_vector)
```

*** =sct
```{r}
msg = "Đừng thay đổi giá trị `survey_vector` and `factor_survey_vector`."
test_object("survey_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_survey_vector", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
msg <- "Bạn đã sử dụng hàm `summary()` hàm chính xác chưa"
test_output_contains("summary(survey_vector)", incorrect_msg = sprintf(msg, "survey_vector"))
test_output_contains("summary(factor_survey_vector)", incorrect_msg = sprintf(msg, "factor_survey_vector"))
success_msg("Rất tốt! Hãy xem kết quả cuối cùng. Khi bạn gán giá trị `\"Male\"` và `\"Female\"` thành giá trị của factor trong `factor_survey_vector`, bạn cho phép R tính toán và hiển thị số lượng các giá trị trong mỗi biến.")
```


--- type:NormalExercise xp:100 skills:1 key:90ecc160d1ebf2f75bf53f9c3843fc1632bdd0a5
## Xác định giới tính

Trong  `factor_survey_vector`, chúng ta có một factor với hai giá trị: Male và Female. Nhưng làm thế nào R xác định được các giá trị này? Hay nói cách khác, R xác định thứ tự thế nào, Male hay Female? 

*** =instructions

Đọc các câu lệnh trong editor và clcik 'Submit Answer' để xem giá trị nào, Male hay Female có giá trị hơn.

*** =hint
Ấn vào nút 'Submit Answer' để xem kết quả được hiển thị trên màn hình tương tác.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Tạo véc-tơ factor_survey_vector 
survey_vector <- c("M", "F", "F", "M", "M")
factor_survey_vector <- factor(survey_vector)
levels(factor_survey_vector) <- c("Female", "Male")

# Male
male <- factor_survey_vector[1]

# Female
female <- factor_survey_vector[2]

# Đánh giá các giá trị của giới tính: Male hay  female  'có giá trị hơn'?
male > female
```

*** =solution
```{r}
# Tạo véc-tơ factor_survey_vector 
survey_vector <- c("M", "F", "F", "M", "M")
factor_survey_vector <- factor(survey_vector)
levels(factor_survey_vector) <- c("Female", "Male")

# Male
male <- factor_survey_vector[1]

# Female
female <- factor_survey_vector[2]

# Đánh giá các giá trị của giới tính: Male hay  female  'có giá trị hơn'?
male > female
```

*** =sct
```{r}
msg = "Không thay đổi các câu lệnh; chỉ ấn 'Submit Answer' và xem kết quả."
test_object("survey_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_survey_vector", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)
test_object("male", undefined_msg = msg, incorrect_msg = msg)
test_object("female", undefined_msg = msg, incorrect_msg = msg)
test_output_contains("male > female", incorrect_msg = msg)
success_msg("Hừm, hình như R không phân biệt giới tính. Hoặc ta cần thảo luận thêm vấn đề này sau vậy ;-).")
```


--- type:NormalExercise xp:100 skills:1 key:9ab0928916bf84ab225713a9a1ce40d9e322c6a0
## Factor có phân biệt thứ tự (1)

Bời vì `"Male"` và `"Female"` có giá trịnh không phân biệt thứ tự, nên R lưu ý với bạn rằng, giá trị lớn hơn không có giá trị trong trường hợp này. Như ta đã thấy trước đó, R gán các giá trị bằng nhau cho factor


Tuy nhiên, không phải lúc nào cũng vậy! Đôi lúc bạn cần phải xử lý fator có phân biệt thứ tự. Trong trường hợp này, ta cần chắc chắn rằng ta đã khai báo các thông tin này với R.

Giả sử rằng bạn đang hướng dẫn một nhóm nhiên cứu gồm có 5 chuyên viên phân tích dữ liệu, và rằng bạn muốn đánh giá quá trình nghiên cứu của họ. Để làm điều này, bạn cần đánh giá tốc độ làm việc của họ thành các giá trị `"slow"`, `"fast"` hay `"insane"`, và lưu kết quả trong véc-tơ `speed_vector`.

*** =instructions
Ở bước đầu tiên, ta cần gán giá trị `speed_vector`, biết rằng

- Chuyên viên một làm nhanh (fast)
- Chuyên viên hai làm chậm (slow)
- Chuyên viên ba làm chậm (slow)
- Chuyên viên bốn làm nhanh (fast)
- Chuyên viên 5 làm rất tệ (insane)

Ta chưa cần xác định các factor này ngay

*** =hint
Gán véc-tơ `speed_vector` một véc-tơ có giá trị, `"fast"`, `"slow"` ...

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Tạo véc-tơ speed_vector
speed_vector <-
```

*** =solution
```{r}
# Tạo véc-tơ speed_vector
speed_vector <- c("fast", "slow", "slow", "fast", "insane")
```

*** =sct
```{r}
test_object("speed_vector",
            incorrect_msg = "Đảm bảo rằng bạn đã gán các giá trị chính xác cho véc-tơ speed_vector. Không sử dụng các chữ in hoa!")
success_msg("Làm rất tốt! Làm bài tập tiếp theo nào.")
```


--- type:NormalExercise xp:100 skills:1 key:279077d10248ce03d5f972939ef8576430a16683
## Factor có phân biệt thứ tự (2)

`speed_vector`  cần được chuyển thành dạng factor có phân biệt thứ tự. Hàm [`factor()`](http://www.rdocumentation.org/packages/base/functions/factor) chuyển đổi `speed_vector` thành factor không phân biệt thứ tự. Để tạo một véc-tơ có thứ tự, bạn cần phải thêm vào hai tham số `ordered` và `levels`.

```
factor(some_vector,
       ordered = TRUE,
       levels = c("lev1", "lev2" ...))
```

Bằng cách đặt tham số `ordered` thành `TRUE` trong hàm [`factor()`](http://www.rdocumentation.org/packages/base/functions/factor), bạn đã chỉ ra rằng các giá trị trong factor được sắp xếp thứ tự. Tham số `levels`  cho phép các giá trị của factor được sắp xếp theo đúng thứ tự.

*** =instructions
Từ véc-tơ `speed_vector`, tạo một véc-tơ có phân biệt thứ tự : `factor_speed_vector`. Đặt tham số `ordered` thành `TRUE`, và đặt `levels` thành `c("slow", "fast", "insane")`.

*** =hint
Sử dụng hàm  [`factor()`](http://www.rdocumentation.org/packages/base/functions/factor) để tạo `factor_speed_vector` dựa trên véc-tơ `speed_character_vector`. Tham số  `ordered` cần được đặt thành `TRUE` bởi lẽ các giá trị này cần được sắp xếp theo thứ tự. Thêm vào đó, ta cần đặt  `levels = c("slow", "fast", "insane")`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Tạo speed_vector
speed_vector <- c("fast", "slow", "slow", "fast", "insane")

# Chuyển véc-tơ speed_vector thành véc-tơ có phân biệt thứ tự
factor_speed_vector <-

# Hiển thị kết quả factor_speed_vector
factor_speed_vector
summary(factor_speed_vector)
```

*** =solution
```{r}
# Tạo speed_vector
speed_vector <- c("fast", "slow", "slow", "fast", "insane")

# Chuyển véc-tơ speed_vector thành véc-tơ có phân biệt thứ tự
factor_speed_vector <- factor(speed_vector, ordered = TRUE, levels = c("slow", "fast", "insane"))

# Hiển thị kết quả factor_speed_vector
factor_speed_vector
summary(factor_speed_vector)
```

*** =sct
```{r}
msg = "Không thay đổi các giá trị ảnh hưởng đến biến `speed_vector`."
test_object("speed_vector", undefined_msg = msg, incorrect_msg = msg)
test_function("factor", args = c("x", "ordered", "levels"),
              incorrect_msg = c("Tham số đầu tiên cần xác định cho  `factor()` là `speed_vector`.",
                                "Bạn cần chắc chắn rằng đã đặt `ordered = TRUE` bên trong hàm `factor()`.",
                                "Bạn cần chắc chắn rằng đã đặt `levels = c(\"slow\", \"fast\", \"insane\")` bên trong hàm `factor()`."))
test_object("factor_speed_vector", eq_condition = "equal",
            incorrect_msg = "Có điểm gì đó chưa chính xác với véc-tơ `factor_speed_vector`; kiểm tra lại `speed_vector`, `ordered = TRUE` and `levels = c(\"slow\", \"fast\", \"insane\")` trong giá trị của hàm `factor()`.")
success_msg("Rất tốt! Hãy xem màn hình tương tác console. Bây giờ đã hiển thị thứ tự của véc-tơ với dấu `<`. Tiếp tục làm bài tập tiếp theo nhé")
```


--- type:NormalExercise xp:100 skills:1 key:db16e69805625bcfde227743a8cbc985f8482a37
## So sánh thứ tự của factor

Sau một ngày làm việc không hiệu quả, 'Chuyên viên phân tích số hai' tiến vào văn phòng làm việc của bạn, bắt đầu phàn nàn rằng 'Chuyên viên số năm' đang làm chậm tiến độ của dự án. Vốn biết rằng chuyên viên số hai có tiếng là láu cá, bạn quyết định kiểm tra thông tin này


Véc-tơ `factor_speed_vector` , giờ đã được sắp xếp thứ tự, cho phép chúng ta so sánh các thành phần khác nhau (chuyên viên phân tích dữ liệu trong trường hợp này). 

*** =instructions
- Sử dụng `[2]` để lựa chọn factor từ véc-tơ `factor_speed_vector`, giá trị này được gán cho chuyên viên hai. Lưu giá trị này thành  `da2`.
- Sử dụng `[5]` để lựa chọn factor từ véc-tơ `factor_speed_vector`, giá trị này được gán cho chuyên viên hai. Lưu giá trị này thành  `d52`.
- Kiểm tra xem `da2` có lớn hơn giá trị `da5` hay không; hiển thị kết quả ra ngoài màn hình. Nhớ rằng bạn có thể sử dụng dấu lớn hơn  `>` trong trường hợp này

*** =hint
- Để lựa chọn giá trị factor của chuyên viên số ba, bạn sẽ cần sử dụng câu lệnh `factor_speed_vector[3]`.
- Để so sánh hai giá trị, bạn có thể sử dụng dấu lớn hơn `>`. Ví dụ: `da3 > da4`.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Tạo véc-tơ factor_speed_vector
speed_vector <- c("fast", "slow", "slow", "fast", "insane")
factor_speed_vector <- factor(speed_vector, ordered = TRUE, levels = c("slow", "fast", "insane"))

# Giá trị factor của chuyên viên số 2
da2 <-

# Giá trị factor của chuyên viên số 5
da5 <-

# Chuyên viên 2 có làm việc nhanh hơn chuyên viên 5?

```

*** =solution
```{r}
# Tạo véc-tơ factor_speed_vector
speed_vector <- c("fast", "slow", "slow", "fast", "insane")
factor_speed_vector <- factor(speed_vector, ordered = TRUE, levels = c("slow", "fast", "insane"))

# Giá trị factor của chuyên viên số 2
da2 <- factor_speed_vector[2]

# Giá trị factor của chuyên viên số 5
da5 <- factor_speed_vector[5]

# Chuyên viên 2 có làm việc nhanh hơn chuyên viên 5?
da2 > da5
```

*** =sct
```{r}
msg = "Đừng thay đổi giá trị của véc-tơ `speed_vector` và `factor_speed_vector`!"
test_object("speed_vector", undefined_msg = msg, incorrect_msg = msg)
test_object("factor_speed_vector", eq_condition = "equal", undefined_msg = msg, incorrect_msg = msg)

msg <- "Bạn đã lấy đúng giá trị của %s chuyên viên chưa? Bạn có thể sử dụng hàm `factor_speed_vector[%s]`."
test_object("da2", eq_condition = "equal", incorrect_msg = sprintf(msg, "second", "2"))
test_object("da5", eq_condition = "equal", incorrect_msg = sprintf(msg, "fifth", "5"))
test_output_contains("da2 > da5", incorrect_msg = "Bạn đã so sánh đúng giá trị của `da2` và `da5` chưa? Bạn có thể sử dụng dấu lớn hơn `>`. Hiển thị kết quả ngoài màn hình")

success_msg("Tuyệt vời! Kết quả trên cho bạn thấy điều gì? Chuyên viên hai phàn nàn về chuyên viên năm trong khi chính chuyên viên số hai mới là người làm chậm tiến độ của dự án! Kết quả này kết thúc chương về factor. Với kiến thức nền tảng vững vàng về véc-tơ, ma trận và factor, bạn đã sẵn sàng để tìm hiểu bảng dữ liệu data frame, một cấu trúc dữ liệu vô cùng quan trọng trong R.")
```



